# Assignment 0

This assignment is meant as practice for your subsequent assignments. It will be graded but will **not influence your overall mark**. While the assignment does not influence your mark, you should still take it seriously, as it will provide you feedback on your programming knowledge and the overall complexity of assignments.

In this assignment we will use some of the concepts learned in Lab1. We will use Python to do text manipulation and analysis, using for loops, conditional statements and functions.


## File System Organization & Paths
Files in an operating system are organized in a tree-like pattern of directories (or folders), which may contain files and other directories. We can directly refer to files or folders in our current path by their name. 

If we want to access sub-folders we specify the hierarchy with a slash `/`. We can go up up the hierarchy using `..`. For example, if we want to access all data in a folder called *data* in our current directory the corresponding path is `data/`. If we want to access a specific file called 'dataset.csv' in this folder we would access it via the path `data/dataset.csv`. If the 'data' folder and dataset file would be stored above our current directory we access it with `../data/dataset.csv`.

In Python we can use the [os](https://docs.python.org/3/library/os.html) library to perform operating system actions. We first import the `os` library to have access to all its functions.

```python
import os
```

<div class="alert alert-info" role="alert">
<h1>Exercise 1 </h1>

<p> Specify the path to the books folder below. The `books/` directory is in your current directory. We can specify it as a *string* -- a sequence of characters -- which are are enclosed by quotation marks in python, e.g., `the_name_of_my_cat = "Meowulf"`</p>
</div>

<!-- grade_id: cell-2cca24359550fa82, points: 0, grade: false, solution: true -->
```python
path = "books/"
# YOUR CODE HERE

```

<!-- grade_id: cell-70903f287e06dc8c, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises, assert_not_equal
# The line below tests to make sure the path isn't an empty string. This kind of test is 
# designed to help you check that your answer has the right basic form.
assert_not_equal(path, "")

```

The 'books' folder contains several subfolders. Each subfolder corresponds to one author and contains some of their works as text files. 

<div class="alert alert-info" role="alert">
<h1>Exercise 2</h1>
<p>Write a line of code to calculate the number of author folders in the books folder. Store the number of author folders in the variable num_author_folders.</p>
</div>

### Hints

- The `os` library contains many functions for reading and manipulating files and folders in your operating system. Have a look at `os.listdir()` for this exercise. You can pass a string to this function, e.g., `os.listdir("my_directory_name")`.
- If you call `os.listdir` on the path to the books folder, it will return a list of the subfolders in the 'books' folder. Determining the number of elements in this list gives you the number of authors. You'll need to find out how to compute the length of a list.

<!-- grade_id: cell-4bf970e1c39d00fb, points: 0, grade: false, solution: true -->
```python
num_author_folders =  len(os.listdir("books"))
```

<!-- grade_id: cell-c826b29c07cf8e23, points: 2, grade: true, solution: false -->
```python
assert_equal(num_author_folders, 5)
```

Now we will access all works by the authors. To do so we will iterate through all author folders in the 'books' folder. Then we will list all files in each of those author subfolders.

<div class="alert alert-info" role="alert">
<h1>Exercise 3</h1>
<p>Using Python, create a list containing a list for each author in the books folder. For each author, add the path for each book in the subfolder.</p><p>For example, if there are two authors, "Alice" and "Bob", and Alice wrote "Aardvarks and Archetypes" and Bob wrote "Bothys I have known" and "Bob: A memoir", then you would want to get `[['Aardvarks and Archetypes'],['Bothys I have known','Bob: A memoir']]`.</p>
</div>

### Hints

- Use a `for` loop to iterate through all folders in the 'books' folder.
- Then list all book folders in the authors' folder (using `os.listdir()` again).
- Use `append` to add the list of paths to `all_books`.
- To provide the full path consisting of 'bookpath/authorname/' you will have to concatenate strings. You can do so by adding the path variable with the author name `path+authorname`. Do not forget to add a dash to that string to specify that you want all files *in* that folder.

<!-- grade_id: cell-8304ede06cb5a58a, points: 0, grade: false, solution: true -->
```python
all_books=[]
for author in os.listdir("books"):
    b="books/"+author
    onesbooks=os.listdir(b)
    all_books.append(onesbooks)
    

   
```

<!-- grade_id: cell-5da0d0148e1cd3c7, points: 10, grade: true, solution: false -->
```python

assert_equal(set([book for books in all_books for book in books]),
             set(['The_Power_of_Movement_in_Plants.txt', 'On_the_Origin_of_Species.txt',
              'Treasure_Island.txt',
              'Adventures_of_Sherlock_Holmes.txt',   'Memoirs_of_Sherlock_Holmes.txt',   'The_Lost_World.txt',
              'A_Treatise_of_Human_Nature.txt',   'An_Enquiry_Concerning_Human_Understanding.txt', 
              'Ivanhoe:_A_Romance.txt', 'The_Lady_of_the_Lake.txt']))
```

We will now calculate how many books there are in total. We will loop through the `all_books` list that you just created. We will increment a counter (a variable that starts at zero, which we use to store the count so far) for the total number of books by adding the number of books in each of the entries of `all_books`.

<div class="alert alert-info" role="alert">
<h1>Exercise 4</h1>
<p>Using Python, calculate the total number of books.</p>
</div>

### Hints

- Use a `for` loop to iterate through `all_books` list that you just created.
- Remember that each list contains another list that contains the book paths.
- The length of this list corresponds to the number of books by each author.
- Increment the counter `book_count` by the length of each book list.

<!-- grade_id: cell-ee04fddb559cd5a7, points: 0, grade: false, solution: true -->
```python
book_count = 0
for book_list in os.listdir("books"):
      b="books/"+book_list
    
      book_count += len(os.listdir(b))


```

<!-- grade_id: cell-f07102023c088fba, points: 8, grade: true, solution: false -->
```python
assert_equal(book_count,10)
```

<div class="alert alert-info" role="alert">
<h1>Exercise 5</h1>
<p>Which author has the most books in our books folder? How many books are there for that author?</p>
</div>

### Hints

- As in Exercise 4 we need to loop through all options and iteratively update the authors' name with the most books (`author_most_books`), as well as the number that author has (`most_books`).
- Since the `all_books` list does not contain the authors' name we need to list the author names again using `os.listdir()`.
- For each author determine the number of books in the folder, as in Exercises 3 and 4. If that author has more books than the current largest number of books stored in `most_books`, update `most_books`, as well as the `author_most_books` variable.

<!-- grade_id: cell-27814d82bb53f8be, points: 0, grade: false, solution: true -->
```python
author_most_books = ""
most_books = 0
for author in os.listdir(path):
        b="books/"+author
        number_of_books = len(os.listdir(b))
    
        if number_of_books > most_books:
            author_most_books = author
            most_books = number_of_books
```

<!-- grade_id: cell-d444aa5641061098, points: 15, grade: true, solution: false -->
```python
assert_equal(author_most_books, "Conan_Doyle")
assert_equal(most_books, 3)
```

<div class="alert alert-info" role="alert">
<h1>Exercise 6</h1>
<p>You might have noticed that the formatting of the author is a bit odd. Let's replace the underscores in the authors' name with spaces.
</p>
</div>

### Hints

- Use 'str.replace()' to change the authors name

<!-- grade_id: cell-6cda115e237eef96, points: 0, grade: false, solution: true -->
```python
author_most_books = author_most_books.replace('_',' ')
```

<!-- grade_id: cell-d340e1ddf438bd98, points: 1, grade: true, solution: false -->
```python
assert_equal(author_most_books, "Conan Doyle")
```

Let's finally have a look at the content of the text files. We will start by looking at Charles Darwin's 'On the Origin of Species' book. In Python we can open files using `open`.  Opening files might sometimes produce errors, so it is good practice to take care of potential errors by wrapping the opening command in a `with` statement (For more details, see the [documentation](https://docs.python.org/3/reference/compound_stmts.html#the-with-statement)).

```python
origin_of_species_path = "./books/Charles_Darwin/On_the_Origin_of_Species.txt"
```

```python
with open(origin_of_species_path) as file:
    origin_of_species = file.readlines()
```

The cell above reads in the file specified in the path `origin_of_species_path` into a variable called `origin_of_species`. Inspect this variable to see what it looks like. 

<div class="alert alert-info" role="alert">
<h1>Exercise 7</h1>
<p>Access the title of On the Origin of Species in the variable (lines 30 to 33).</p>
</p>
</div>

## Hints

 - What is the type of the variable `origin_of_species`? Remember how we can access elements in this data type (see the content from lab 1).

<!-- grade_id: cell-5d9541b0b9bd25f5, points: 0, grade: false, solution: true -->
```python
origin_of_species_title = (origin_of_species)[30:33]


```

<!-- grade_id: cell-3eead2e4b04842c1, points: 1, grade: true, solution: false -->
```python
assert_equal(origin_of_species_title, ['ON THE ORIGIN OF SPECIES.\n',  '\n',  'OR THE PRESERVATION OF FAVOURED RACES IN THE STRUGGLE FOR LIFE.\n'])
```

Looking at word occurrences and co-occurrences is a an important technique in texts analysis. We are often not really interested in differentiating words by their punctuation or whether they are capitalized. Therefore, it is important to format words consistently (for example by having all words in lowercase and by removing punctuation).

We will start by creating a function to format each word in the texts. Python provides a string containing all punctuation characters. We can access it by importing the `string` type and creating a set of all punctuation characters with `set(string.punctuation)`.

```python
import string
punctuation = set(string.punctuation)
```

Now we will specify our formatting function. The function takes a word as input parameter and returns the formatted word. 

<div class="alert alert-info" role="alert">
<h1>Exercise 8</h1>
<p>Implement the formatting function. The function should remove punctuation characters and make sure that all characters in the word are lowercase.</p>
</p>
</div>

## Hints
- You can loop through characters in a word using a `for` loop.
- You can check if a character is in the punctuation set using a `if` clause together with `in`.
- You should append all characters that are not punctuation characters to the list `non_punctuation_chars`.
- Then you need to join the characters to a string. The easiest way to do so is to use `"".join(list_of_chars)`, i.e. we concatenate the characters in the list with the empty string `""`.
- Convert the word to a lowercase version using `str.lower()` and store that version in `new_word_lower`
- Finally, don't forget to return `new_word_lower`!

<!-- grade_id: cell-6cbcbaae8985bb88, points: 0, grade: false, solution: true -->
```python
def format_word(word):
    non_punctuation_chars = []
    for character in word:
        if character not in punctuation:
            non_punctuation_chars.append(character)
    new_word = "".join(non_punctuation_chars)
    new_word_lowercase =str .lower(new_word)
    
   
    return (new_word_lowercase)


```

<!-- grade_id: cell-13b33e033b2e16dc, points: 15, grade: true, solution: false -->
```python
assert_equal(format_word("TESTWORD"), 'testword')
assert_equal(format_word(""), '')
assert_equal(format_word("!&*"), '')
assert_equal(format_word("ARGH!&*"),'argh')
assert_raises(TypeError,format_word, 12)
```

We will write a second function that will count how many times a word occurs in a book file.
 
<div class="alert alert-info" role="alert">
<h1>Exercise 9</h1>
<p>Implement the function word_occurence_in_book. The function takes two arguments, the path of the text file and a word. It returns the number of times the (formatted) word occurs in the list.</p>
</div>

### Hints
- First open the file (as we did for Exercise 7) and read each line in the file.
- To get individual words from a line we can use `str.split(" ")`. This method takes a string (a line in the text) and splits it into individual strings at each whitespace (" "). It returns a list of strings.
- Iterate through the list of words in each line. Remember to format the word using the function you just implemented and store the result in `formatted_word`. Then check if the formatted word is equal to the word we are looking for `count_word` (use an `if` statement). If it is equal you should increment the count.

<!-- grade_id: cell-f386c64f861fbd39, points: 0, grade: false, solution: true -->
```python
def word_occurence_in_book(book_path, count_word):
    count = 0
    with open(book_path) as file:
         book_lines = file.readlines()
    for line in book_lines:
        words = line.split(" ")
        for word in words:
            formatted_word = format_word(word)
            if count_word==formatted_word:
                      count +=1
    return count

```

Output:
```
9471
```

<!-- grade_id: cell-a2c5832ea9de834b, points: 25, grade: true, solution: false -->
```python
assert_equal(word_occurence_in_book("books/Charles_Darwin/On_the_Origin_of_Species.txt", "the"), 9471)
assert_equal(word_occurence_in_book("books/Charles_Darwin/The_Power_of_Movement_in_Plants.txt", "book"), 1)
assert_equal(word_occurence_in_book("books/Charles_Darwin/The_Power_of_Movement_in_Plants.txt", "move"), 110)
assert_equal(word_occurence_in_book("books/David_Hume/An_Enquiry_Concerning_Human_Understanding.txt", "move"), 8)
assert_equal(word_occurence_in_book("books/Louis_Stevenson/Treasure_Island.txt", "pirate"), 6)
assert_raises(FileNotFoundError, word_occurence_in_book, "", "move")
```

Output:
```
/opt/conda/lib/python3.6/unittest/case.py in assertEqual(self, first, second, msg)
    827         """
    828         assertion_func = self._getAssertEqualityFunc(first, second)
--> 829         assertion_func(first, second, msg=msg)
    830 
    831     def assertNotEqual(self, first, second, msg=None):

/opt/conda/lib/python3.6/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    820             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    821             msg = self._formatMessage(msg, standardMsg)
--> 822             raise self.failureException(msg)
    823 
    824     def assertEqual(self, first, second, msg=None):

AssertionError: 10435 != 9471
```

Now, we are going to count the occurrence of words for the works of each author. Previously we have been using lists or individual variables to store our results. Since we are looking at the occurrence for each author in the 'book' folder and that folder could potentially contain many authors, a nicer way to store our counts is to use a dictionary. You can read more about dictionaries [here](https://docs.python.org/3/tutorial/datastructures.html#dictionaries).

We will store the author name as the key of the dictionary and its value will correspond to how many times the word occurs in their books, like this: ```{author_name: 35}```.

<div class="alert alert-info" role="alert">
<h1>Exercise 10</h1>
<p>For each author and each book by that author count the occurrence of the word <em>mind</em>. Store your results in a dictionary. Make sure that you count the formatted version of the word!</p>
</div>

<!-- grade_id: cell-42c212087d45a818, points: 0, grade: false, solution: true -->
```python
author_count_dictionary = {}
for author in os.listdir(path):
    author_books = os.listdir(path+author)
    author_count = 0
    for book in author_books:
            book_path = path+author+"/"+book
            count_in_book = word_occurence_in_book(book_path, "mind")
            author_count += count_in_book
            author_count_dictionary[author]=author_count
print(author_count_dictionary)

```

<!-- grade_id: cell-779b15c6421f0999, points: 22, grade: true, solution: false -->
```python
assert_equal(author_count_dictionary, {'Charles_Darwin': 55,  'Conan_Doyle': 103,  'David_Hume': 667,  'Louis_Stevenson': 36,  'Walter_Scott': 56})
```